Przed oddaniem zadania upewnij się, że wszystko działa poprawnie.
**Uruchom ponownie kernel** (z paska menu: Kernel$\rightarrow$Restart) a następnie
**wykonaj wszystkie komórki** (z paska menu: Cell$\rightarrow$Run All).

Upewnij się, że wypełniłeś wszystkie pola `TU WPISZ KOD` lub `TU WPISZ ODPOWIEDŹ`, oraz
że podałeś swoje imię i nazwisko poniżej:

```python
NAME = "Marcel Cielinski"
```

---

# Środowisko w laboratorium

W tym notebooku znajdują się zadania, mające na celu zapoznanie się ze środowiskiem, w którym będzie się odbywać laboratorium. Wykonaj wszystkie zadania zgodnie z zamieszczonymi przy nich instrukcjami.


W laboratorium wykorzystywane są zeszyty Jupyter Notebook. Są to interaktywne dokumenty mogące zawierać kod oraz wyniki jego działania, wraz z dodatkową treścią, dodawaną przy użyciu języka [Markdown](https://www.markdownguide.org/). W zeszytach Jupyter Notebook można wykorzystywać wiele języków: Python, R, Julia itp.; na laboratorium wykorzystywany będzie język Python.
    
Przygotowanie środowiska polega na zainstalowaniu wymaganych paczek z dostarczonego pliku `requirements.txt`:

```console
pip install -r requirements.txt
```

Polecamy wykorzystanie [wirtualnych środowisk](https://packaging.python.org/guides/installing-using-pip-and-virtual-environments/) w celu izolacji od systemowego interpretera Python (plik `requirements.txt` może się zmieniać na kolejnych laboratoriach).

Po zainstalowaniu paczek środowiska należy uruchomić serwer Jupyter poleceniem:

```console
jupyter notebook
```

Po uruchomieniu w przeglądarce otworzy się strona Jupyter, z której można wybrać plik zeszytu (o rozszerzeniu `*.ipynb`), aby go otworzyć. Jeżeli strona nie otworzyła się automatycznie, można wykorzystać link, który pojawi się w konsoli.


Zeszyt może zawiera komórki dwóch głównych rodzajów. Domyślnie tworzone są komórki zawierające kod (`Code` na pasku narzędzi). Pozwalają one na wprowadzanie i wykonywanie kodu (tu: w języku Python). Aby utworzyć nową komórkę poniżej aktualnej wybierz symbol <kbd>+</kbd> z paska narzędzi.

Zeszyty jupyterowe umożliwiają również wywoływanie komend w konsoli. W tym celu należy dołączyć z przodu znak `!`:

```python
!echo "This runs in shell"
```

Typ komórki można zmienić wybierając ją, a następnie wybierając z listy na pasku narzędzi np. `Markdown`.

Każdą komórkę można wykonać, wybierając symbol <kbd>▷</kbd> lub wciskając klawisze <kbd>⇧ Shift</kbd>+<kbd>⏎ Enter</kbd>. W każdej chwili możesz zatrzymać wykonywanie komórki wciskając symbol <kbd>⏹️</kbd>.

Wartym odnotowania jest fakt, że wszystkie komórki w ramach jednego zeszytu współdzielą te same zmienne i importy. Istotna jest tutaj kolejność wywoływania komórek - ostatnia wykonana komórka nadpisuje pokrywające się zmienne. Przestrzeń zmiennych i importów można wyczyścić restartując kernel (z paska menu: <kbd>Kernel</kbd>-><kbd>Restart</kbd>).

Narzędzie jest proste w obsłudze: polecamy przejrzeć pasek menu oraz pasek narzędzi i sprawdzić dostępne opcje.


### Zadanie 0.1

Utwórz pod niniejszą komórką dwie komórki: jedną typu `Code`, a drugą typu `Markdown`, wypełnij je dowolną treścią i wykonaj.


```python
a = 2
b = 3

f'{a} + {b} = {a+b}'
```

Output:
```
'2 + 3 = 5'
```

## Dowolna treść
### Dowolna treść
Dowolna treść

Poniżej znajdują się przykładowe zadania. Wykonaj je, wypełniając pola 'TU WPISZ KOD' oraz 'TU WPISZ ODPOWIEDŹ'. Do zadania 0.3 zostały dołączone przykładowe testy sprawdzające działanie kodu - nie są to jednak wszystkie testy przygotowane w laboratorium.


### Zadanie 0.2

Zapisz wzór średniej harmonicznej, używając składni $\LaTeX$.

Przykład: `$$y=ax+b$$`
$$y=ax+b$$


<!-- grade_id: harmonic_mean_eq, points: 0, grade: true, solution: true -->
TU WPISZ ODPOWIEDŹ

$$x_{HM}=\frac {N}{{\frac {1}{x_{1}}}+{\frac {1}{x_{2}}}+\cdots +{\frac {1}{x_{N}}}}$$
$$x_{HM}=\frac {N}{\sum \limits _{i=1}^{N}{\frac {1}{x_{i}}}}$$
$$x_{HM}=\left(\frac {\sum \limits _{i=1}^{N}x_{i}^{-1}}{N}\right)^{-1}$$

### Zadanie 0.3

Zaimplementuj średnią harmoniczną w *czystym* pythonie. Funkcja `harmonic_mean` ma przyjmować listę dodatnich liczb i zwracać średnią harmoniczną.

Ważne: obsłuż błędne przypadki wywołując błąd `ValueError`.


<!-- grade_id: harmonic_mean_code, points: 0, grade: false, solution: true -->
```python
def harmonic_mean(data):
    """Return harmonic mean for given list of numbers."""
    # TU WPISZ KOD
    # raise NotImplementedError()

    # try:
    #     result = len(data) / sum(1. / val for val in data)
    # except ZeroDivisionError:
    #     raise ValueError

    hm_sum = 0
    for val in data:
        if val <= 0:
            raise ValueError
        hm_sum += 1./val
    
    try:
        result = len(data) / hm_sum
    except ZeroDivisionError:
         raise ValueError

    return result
```

<!-- grade_id: harmonic_mean_tests, points: 0, grade: true, solution: false -->
```python
# Sprawdzenie działania funkcji
from nose.tools import assert_almost_equal, assert_raises

assert_almost_equal(harmonic_mean([2, 3, 5, 7, 8]), 3.843, places=3)
assert_almost_equal(harmonic_mean([2.5, 5.23, 8.64, 9.87, 23.5]), 5.877, places=3)
assert_raises(ValueError, harmonic_mean, [-1, 2])
assert_raises(ValueError, harmonic_mean, [])

```

```python

```